---
title_meta  : 第二章
title       : 欄位相關的技巧
description : 實務上我們很常會有新增變數、刪除變數或者重新命名變數...等等的需求，在資料框的結構中，其實就是針對欄位進行整理，我們將在本章節學習這些技巧，一場爭奪 One Piece 的海上冒險故事！

--- type:NormalExercise lang:r xp:100 skills:4
## 新增欄位

草帽海賊團的廚師賓什莫克·香吉士要為船員們準備餐點，卻發現主要角色設定遺漏了大家最喜愛的料理，讓我們來幫助香吉士，將最喜愛的料理 `favorite_food` 加入 `straw_hat_df` 中。

*** =instructions
- 將右邊編輯區已經定義好的 `favorite_food` 向量加入 `straw_hat_df` 中。

*** =hint
- 在編輯區鍵入 `straw_hat_df$favorite_food <- favorite_food`

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# 最喜愛料理的向量
favorite_food <- c("肉", "搭配酒的食物", "橘子", "魚", "搭配紅茶的食物", "甜食", "搭配咖啡的食物", "搭配可樂的食物", "牛奶")

# 將向量加入資料框中成為新的欄位
straw_hat_df$favorite_food <- 
```

*** =solution
```{r}
# 最喜愛料理的向量
favorite_food <- c("肉", "搭配酒的食物", "橘子", "魚", "搭配紅茶的食物", "甜食", "搭配咖啡的食物", "搭配可樂的食物", "牛奶")

# 將向量加入資料框中成為新的欄位
straw_hat_df$favorite_food <- favorite_food
```

*** =sct
```{r}
test_object(favorite_food,
            undefined_msg = "&#19981;&#38656;&#35201;&#21034;&#38500;&#21407;&#26412;&#24171;&#20320;&#23450;&#32681;&#22909;&#30340;&#21521;&#37327;&#65281;", 
            incorrect_msg = "&#19981;&#38656;&#35201;&#21034;&#38500;&#21407;&#26412;&#24171;&#20320;&#23450;&#32681;&#22909;&#30340;&#21521;&#37327;&#65281;") 

test_data_frame("straw_hat_df", columns = "favorite_food",
                undefined_cols_msg = "&#27298;&#26597;&#26159;&#21542;&#26377;&#23559;&#24050;&#32147;&#23450;&#32681;&#22909;&#30340; `favorite_food` &#21521;&#37327;&#21152;&#20837; `straw_hat_df` &#20013;&#65311;",
                incorrect_msg = "&#27298;&#26597;&#26159;&#21542;&#26377;&#23559;&#24050;&#32147;&#23450;&#32681;&#22909;&#30340; `favorite_food` &#21521;&#37327;&#21152;&#20837; `straw_hat_df` &#20013;&#65311;")

success_msg("&#20320;&#24171;&#20102;&#39321;&#21513;&#22763;&#19968;&#20491;&#22823;&#24537;&#65292;&#36889;&#27171;&#20182;&#25165;&#30693;&#36947;&#35442;&#28310;&#20633;&#21738;&#20123;&#26009;&#29702;&#65292;&#22826;&#26834;&#20102;&#65281;")
```

